# Exercise 1

An easy one to start with. Calculate the number of letters in your name:

<!-- grade_id: get_length_name, points: 0, grade: false, solution: true -->
```python
def get_length_name(first_name, last_name):
    length = 0
    
    ### BEGIN SOLUTION    

    ### END SOLUTION
    
    return length
```

```python
first_name = "Ada"
last_name = "Lovelace"

get_length_name(first_name, last_name)
```

Output:
```
0
```

Verify your answer (you typically get 1 answer for free, so you are aware of how the exercises are validated):

<!-- grade_id: get_length_name, points: 5, grade: true, solution: false -->
```python
assert get_length_name('Grace','Hopper') == 11

### BEGIN HIDDEN TESTS


assert get_length_name('Anita','Borg') == 9


assert get_length_name('','') == 0


assert get_length_name('f','') == 1


assert get_length_name('','jj') == 2
### END HIDDEN TESTS
```

# Exercise 2

The second exercise focuses on arithmetic and arrays.
Find the sum of the elements of two arrays.

<!-- grade_id: get_array_sums, points: 0, grade: false, solution: true -->
```python
arr1 = [1, 2, 3, 4, 5]
arr2 = [6, 7, 8, 9, 10]

def get_array_sums(arr1, arr2):
    sum1 = 0  # or other ways you prefer
    sum2 = 0
    
    ### BEGIN SOLUTION 

    
    ### END SOLUTION
    
    return sum1+sum2
```

Verify your answer:

<!-- grade_id: get_array_sums, points: 5, grade: true, solution: false -->
```python
assert get_array_sums([1, 2, 3, 4, 5],[6, 7, 8, 9, 10]) == 55

### BEGIN HIDDEN TESTS


assert get_array_sums([1, 2, 3, 4, 5],[6, 7, 8, 9, 10]) == 55


assert get_array_sums([],[6, 7, 8, 9, 10]) == 40


assert get_array_sums([1, 2, 3, 4, 5],[]) == 15


assert get_array_sums([],[]) == 0
### END HIDDEN TESTS
```

# Exercise 3

Calculate how many years you need to obtain a certain amount by starting with an initial amount and a given interest rate. Don't forget: we use compound interest: you get interest on your interest.

<!-- grade_id: calculate_interest, points: 0, grade: false, solution: true -->
```python
amount = 1000
target = 2000
interest_rate = 0.02

def calculate_interest(amount, target, interest_rate):
    number_of_years = 0  # or other ways you prefer
    
    ### BEGIN SOLUTION    

    ### END SOLUTION
    
    return number_of_years
```

Verify your answer:

<!-- grade_id: calculate_interest, points: 8, grade: true, solution: false -->
```python
assert calculate_interest(1000,2000,0.02) == 36

### BEGIN HIDDEN TESTS


assert calculate_interest(1000,1500,0.02) == 21


assert calculate_interest(1,1500,0.8) == 13


assert calculate_interest(1,80000000000,0.01) == 2524
### END HIDDEN TESTS
```

# Exercise 4

Remove the duplicates from a list.

<!-- grade_id: remove_duplicates, points: 0, grade: false, solution: true -->
```python
input_list = ["Anita", "Ada", "Anita", "Grace", "Grace"]

def remove_duplicates(input_list):
    output_list = [] # or other ways you prefer
    
    ### BEGIN SOLUTION    

    ### END SOLUTION
    
    return output_list
```

Verify your answer:

<!-- grade_id: remove_duplicates, points: 4, grade: true, solution: false -->
```python
assert remove_duplicates(["Anita", "Ada", "Anita", "Grace", "Grace"]) == ['Anita', 'Ada', 'Grace']

### BEGIN HIDDEN TESTS


assert remove_duplicates([]) == []


assert remove_duplicates(["Anita", "Ada", "Anita"]) == ['Anita', 'Ada']


assert remove_duplicates(["123", "345", "345"]) == ["123", "345"]
### END HIDDEN TESTS
```

# Exercise 5

Find the total number of hours taught by a particular individual. The input is two dictionaries: one containing the courses and their organisers, and one containing the courses and hours of lectures.

<!-- grade_id: cell-2aa20234cf44249b, points: 0, grade: false, solution: true -->
```python
courses = {"Analytics": "Xuefei", "Mega analytics" : "Xuefei", "Basic analytics" : "Yannis"}
hours = {"Analytics": 15, "Mega analytics": 10, "Basic analytics":5}
organizer = "Xuefei"

def find_hours_for_organizer(organizer, courses, points):
    total = 0 # or other ways you prefer
    
    ### BEGIN SOLUTION    

    ### END SOLUTION
    
    return total
```

Verify your answer:

<!-- grade_id: find_points_for-organizer, points: 6, grade: true, solution: false -->
```python
courses = {"Analytics": "Xuefei", "Mega analytics" : "Xuefei", "Basic analytics" : "Yannis"}
hours = {"Analytics": 15, "Mega analytics": 10, "Basic analytics":5}

assert find_hours_for_organizer("Xuefei", courses, points) == 25

### BEGIN HIDDEN TESTS


assert find_hours_for_organizer("Yannis", courses, points) == 5


assert find_hours_for_organizer("", courses, points) == 0
### END HIDDEN TESTS
```