# Exercises with libraries

Before you begin, load the wine dataset, then perform the tasks outlined below.

```python
from sklearn import datasets as ds
import pandas as pd
import numpy as np

dataset = ds.load_wine()
```

First, retrieve all the names of the independent variables:

<!-- grade_id: get_variable_names, points: 0, grade: false, solution: true -->
```python
def get_variable_names():
    variables = []
    
    ### BEGIN SOLUTION    

    
    ### END SOLUTION
    
    return variables
```

Your answer will be verified below (no need for you to do anything).

<!-- grade_id: get_variables_names, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS


assert get_variable_names() == ['alcohol',
 'malic_acid',
 'ash',
 'alcalinity_of_ash',
 'magnesium',
 'total_phenols',
 'flavanoids',
 'nonflavanoid_phenols',
 'proanthocyanins',
 'color_intensity',
 'hue',
 'od280/od315_of_diluted_wines',
 'proline']
### END HIDDEN TESTS
```

Now, create a function that calculates the mean and median of a given variable in the dataset (enter the variable's name as a string):

<!-- grade_id: calculate_mean_median, points: 0, grade: false, solution: true -->
```python
def calculate_mean_median(variable):
    mean = 0
    median = 0
    
    ### BEGIN SOLUTION    
    # First, import the dataframe using the raw data (values), and the features
      
    # This makes it easy to quickly obtain the data of a variable using df[variable]
    # We use NumPy to calculate the mean and median

    ### END SOLUTION
    
    return mean, median
```

Finally, verify your answer (note: the function allclose is used to avoid any difference in rounding):

<!-- grade_id: calculate_mean_median, points: 4, grade: true, solution: false -->
```python
assert np.allclose(calculate_mean_median('flavanoids'),(2.0292696629213474, 2.135))

### BEGIN HIDDEN TESTS


assert np.allclose(calculate_mean_median('total_phenols'),(2.295112359550562, 2.355))


assert np.allclose(calculate_mean_median('od280/od315_of_diluted_wines'),(2.6116853932584254, 2.78))


assert np.allclose(calculate_mean_median('ash'),(2.3665168539325854, 2.36))
### END HIDDEN TESTS
```