# Calculating summary statistics yourself

You are now tasked with performing some calculations yourself.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
import sklearn.datasets as dsets
```

## The dataset

```python
# First, we again load one of the datasets from the sklearn.datasets package
dataset = dsets.load_boston()

# You can print the description by outcommenting this line
#print(dataset['DESCR'])

# Let's create a dataframe for the independent variables
housing_data = pd.DataFrame(data=dataset['data'],columns=dataset['feature_names'])

# Now, we create a dataframe for the independent variable
house_prices = pd.DataFrame(data=dataset['target'],columns=['target'])
```

```python
print(housing_data.describe())
```

Output:
```
CRIM          ZN       INDUS        CHAS         NOX          RM  \
count  506.000000  506.000000  506.000000  506.000000  506.000000  506.000000   
mean     3.613524   11.363636   11.136779    0.069170    0.554695    6.284634   
std      8.601545   23.322453    6.860353    0.253994    0.115878    0.702617   
min      0.006320    0.000000    0.460000    0.000000    0.385000    3.561000   
25%      0.082045    0.000000    5.190000    0.000000    0.449000    5.885500   
50%      0.256510    0.000000    9.690000    0.000000    0.538000    6.208500   
75%      3.677083   12.500000   18.100000    0.000000    0.624000    6.623500   
max     88.976200  100.000000   27.740000    1.000000    0.871000    8.780000   

              AGE         DIS         RAD         TAX     PTRATIO           B  \
count  506.000000  506.000000  506.000000  506.000000  506.000000  506.000000   
mean    68.574901    3.795043    9.549407  408.237154   18.455534  356.674032   
std     28.148861    2.105710    8.70725
```

## The interquantile range (ICQ)

Now, it is up to you to calculate the interquantile range of a particular variable (enter the variable's name as a string):

<!-- grade_id: calculate_icq, points: 0, grade: false, solution: true -->
```python
def calculate_icq(variable):
    icq = 0
    
    ### BEGIN SOLUTION    

    ### END SOLUTION
    
    return icq
```

Verify your answer:

<!-- grade_id: calculate_icq, points: 4, grade: true, solution: false -->
```python
assert np.isclose(calculate_icq('INDUS'),12.91)

### BEGIN HIDDEN TESTS


assert np.isclose(calculate_icq('PTRATIO'),2.8000000000000007)


assert np.isclose(calculate_icq('TAX'),387.0)


assert np.isclose(calculate_icq('LSTAT'),10.005000000000003)
### END HIDDEN TESTS
```

## Proportions

Now, convert the frequency table we saw earlier into a table of proportions instead:

<!-- grade_id: calculate_proportions, points: 0, grade: false, solution: true -->
```python
variable = housing_data['CHAS']

def calculate_prop(variable):
    proportions = []
    
    ### BEGIN SOLUTION

    
    ### END SOLUTION
    
    return sorted(proportions)
```

Verify your answer:

<!-- grade_id: calculate_prop, points: 4, grade: true, solution: false -->
```python
variable = housing_data['CHAS']
assert np.allclose(calculate_prop(variable), [0.0691699604743083, 0.9308300395256917])

### BEGIN HIDDEN TESTS

# A different dataset is used to verify your answer as well
dataset = dsets.load_wine()
wine_classes = pd.DataFrame(data=dataset['target'],columns=['target'])
variable = wine_classes['target']
assert np.allclose(calculate_prop(variable), [0.2696629213483146, 0.33146067415730335, 0.398876404494382])
### END HIDDEN TESTS
```